## Lab06: Web Scraping

### Learning Objectives
By the end of this lab, students will be able to:

- **Remember**: Identify key concepts in web scraping (HTTP requests, HTML, tags, attributes).
- **Understand**: Explain how `requests` retrieves web content and how `BeautifulSoup` parses HTML.
- **Apply**: Use `requests` and `BeautifulSoup` to extract targeted information from a webpage.
- **Analyze**: Inspect webpage structure to choose appropriate HTML selectors.
- **Evaluate**: Assess the robustness of selectors and handle missing or inconsistent data.
- **Create**: Produce a clean, readable presentation of scraped data.

### Tasks
- Send an HTTP **GET** request to a webpage.
- Parse the HTML content using `BeautifulSoup`.
- Extract selected elements (e.g., text, links, prices).
- Handle missing data with sensible defaults.
- Display results in a clear, structured format.

### Expected Outcome
A working script that demonstrates the full web scraping workflow—from data retrieval to formatted output—using ethical and responsible practices.

### Notes
- Website structures may vary; always inspect the HTML before writing selectors.
- This lab focuses on **technique and workflow**, not on scraping a specific website.
- Be mindful of website policies and ethical considerations when scraping.

```python
# import library
import requests
from bs4 import BeautifulSoup
import types, importlib, requests, sys, re

import otter
grader = otter.Notebook()

import ssl
ssl._create_default_https_context = ssl._create_unverified_context
```

# Q1: Scrape Title

**Objective:**
Scrape all the titles of the books listed on the Books to Scrape website (http://books.toscrape.com/).

**Instructions:**
- Use requests to send a GET request to the Books to Scrape homepage.
- Parse the HTML content using BeautifulSoup.
- Extract the book titles, which are inside "h3" tags with a nested "a" tag, and print them.

```python
url = "https://books.toscrape.com/"

# Send GET request to the website
response = ...
# Parse the page content using BeautifulSoup
soup = BeautifulSoup(...)
# Find all the title tags and print them
titles = ...
title_list = [title.text.strip() for title in titles]
title_list
```

```python
grader.check("q1")
```

# Q2: Web Scraping (headlines)
**Objective:**
Scrape headlines from the website https://news.ycombinator.com/.

**Task:**
- Send a GET request to the URL using requests.
- Parse the HTML content with BeautifulSoup.
- Extract and print the headlines inside <a> tags with class storylink.
Handle cases where no headlines are found.

```python
# URL of the news site to scrape
url = "https://news.ycombinator.com/"

# Send GET request to the website
response = ...

# Check if the request was successful (status code 200)
if response.status_code == 200:
    print("Request successful!")
else:
    print(f"Request failed with status code: {response.status_code}")
    
# Parse the page content using BeautifulSoup
soup = ...

# Debugging: Print the raw HTML content to see what we are working with
print(soup.prettify()[:1000])  # Print first 1000 characters to inspect the HTML

# Find all the headline links (they are in <a> tags with class 'storylink')
headlines = ...
# Print each headline
title_list = [title.text.strip() for title in headlines]
title_list
```

```python
grader.check("q2")
```

# Q3. Web Scraping (links)
**Objective:**
Scrape all the links (from <a> tags) from the Python homepage (https://www.python.org/).

**Instructions:**
- Use requests to send a GET request to the Wikipedia homepage.
- Parse the HTML content using BeautifulSoup.
- Extract all the links (href attributes) from <a> tags and print them.

```python
url= "https://www.python.org/"

# Send GET request to the website
response = ...

# Parse the page content using BeautifulSoup
soup = ...

# Find all the anchor tags (<a>) which contain links
links = ...
# Print every non-empty href
href_list = [link.get("href") for link in links if link.get("href")]

```

```python
grader.check("q3")
```

# Q4. Web Scraping (quote texts)
**Objective:**
Scrape all the quote texts from the Quotes to Scrape website (https://quotes.toscrape.com/).

**Instructions:**
- Use requests to send a GET request to the Quotes to Scrape homepage.
- Parse the HTML content using BeautifulSoup.
- Extract all the quote texts, which are inside <span> tags with the class 'text', and print them.

```python
# URL of the blog to scrape
url = "https://quotes.toscrape.com/"

# Send GET request to the website
response = ...

# Parse the page content using BeautifulSoup
soup = ..

# Find all the article titles (here, they are in <span> tags with class 'text')
titles = ...
# Print each quote
title_list = [title.text.strip() for title in titles]
title_list
```

```python
grader.check("q4")
```

# Q5. Scraping a Simple Book Store Website for:
    - Book titles
    - Prices
    - Ratings
    - Availability
- URL: http://books.toscrape.com/
- Here, you will work with a scraper example (Book Titles + Prices)

```python
url = "http://books.toscrape.com/"

response = requests.get(url)

if response.status_code == 200:
    print("Request successful!")
else:
    print(f"Request failed with status code: {response.status_code}")

soup = ...

# Find all book containers
books = ...

# Extract title and price
book_data = []
for book in books:
    title = ...
    price = ...
    book_data.append((title, price))

# Show first 5
for title, price in book_data[:5]:
    print(f"{title} — {price}")
```

```python
grader.check("q5")
```

# Q6: Web Scraping newegg for Product Listings

**Objective:** Write a Python script to scrape newegg search results for a specific product (e.g., laptops), extract product details, and display them.

<!-- BEGIN QUESTION -->

### Requirements

- Send an HTTP **GET** request to a laptop search results page using `requests`.
- Parse the HTML response using `BeautifulSoup`.
- Extract each product’s **title** and **price**.
- Display the **first 10 products** in a clean, single-line, aligned format with:
  - a 1-based index (Start counting from 1, not 0)  
  - a truncated product title  
  - the price prefixed by `Price:`
- Use default values (`"No Title"`, `"No Price"`) if any data is missing.

### Example Output

1. Auusda 16" Business Laptop, Windows 11 Pro Laptop Computer, AMD Ryzen 7430U, 32G | Price: $439.99–

```python
url = "https://www.newegg.com/p/pl?d=laptop"


Type your code here...
```

<!-- END QUESTION -->

## Congratulations! You have finished Lab06! 

- Below, you will see a cell. Running this cell will automatically generate a zip file with your autograded answers.

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False)
```

Tests for q1 (points: [1, 1, 1, 2])
```
>>> 'Himalayas' in title_list[-1]
True
```
```
>>> len(title_list) == 20
True
```
```
>>> all((isinstance(t, str) for t in title_list))
True
```
```
>>> sum(('...' in t for t in title_list)) >= 5
True
```


Tests for q2 (points: [3, 3, 3])
```
>>> title_list[-1] == 'More'
True
```
```
>>> len(title_list) == 61
True
```
```
>>> title_list[0] == '1.'
True
```


Tests for q3 (points: [3, 3, 4])
```
>>> expected_slice = ['https://www.python.org/psf/', 'https://docs.python.org', 'https://pypi.org/', '/jobs/', '/community/']
>>> href_list[3:8] == expected_slice
True
```
```
>>> len(href_list) > 0
True
```
```
>>> all((isinstance(href, str) and href.strip() != '' for href in href_list))
True
```


Tests for q4 (points: [3, 3, 3])
```
>>> len(title_list) == 10
True
```
```
>>> all((isinstance(t, str) and t.strip() != '' for t in title_list))
True
```
```
>>> all((quote.startswith('“') and quote.endswith('”') for quote in title_list))
True
```


Tests for q5 (points: [2, 2, 2, 2])
```
>>> len(book_data) == 20
True
```
```
>>> isinstance(book_data, list)
True
```
```
>>> all((isinstance(item, tuple) and len(item) == 2 for item in book_data))
True
```
```
>>> all(('<' not in title and '>' not in title for (title, _) in book_data))
True
```
