# Lab 07 — JSON and Web APIs
### UN5550 — Introduction to Data Science

This lab provides hands-on experience reading nested JSON, requesting data from real web APIs, and organizing responses into pandas DataFrames. Using Pokémon, GitHub, Houghton weather, iTunes, and World Bank data, you will practice data collection, cleaning, visualization, and interpretation.

```python
import json
import requests
import pandas as pd
import numpy as np
import matplotlib.pyplot as plt
from IPython.display import display
import otter
grader = otter.Notebook()
```

### Provided request helper
This helper sends a real GET request, prints its URL/status, checks for an HTTP error, and parses JSON. Use it in the short request cells below. No API key is needed for these examples. A failed request stops its cell; investigate the error before continuing that example.

```python
def get_json(url, params=None):
    response = requests.get(url, params=params, timeout=30)
    print("URL:", response.url, "Status:", response.status_code)
    response.raise_for_status()
    return response.json()
```

## Do Not Modify the JSON Data Below

```python
family_text = '''{
    "name": "Grandma",
    "age": 94,
    "children": [
        {
            "name": "Dad",
            "age": 60,
            "children": [
                {"name": "Me", "age": 24},
                {"name": "Brother", "age": 22}
            ]
        },
        {
            "name": "My Aunt",
            "children": [
                {"name": "Cousin 1", "age": 34},
                {"name": "Cousin 2", "age": 36}
            ]
        }
    ]
}'''
```

## Q1. Read Nested JSON (15 points)

Write a function `family_summary(text)` that accepts a JSON string representing a family tree.

Parse the string using `json.loads()` and return a tuple containing:

1. The name at the top level of the family tree.
2. The name of that person’s first child.
3. The age of the first child of that person’s second child.

Return the values in this exact order:

`(root_name, first_child_name, cousin_age)`

Extract the values from the parsed JSON rather than hard-coding them. Do not modify the supplied JSON.

```python
...
def family_summary(text):
    tree = ...
    root_name = ...
    first_child_name = ...
    cousin_age = ...
    result = ...
    return result
```

```python
grader.check("q1")
```

## Q2. Pokémon API (15 points)

Use the JSON response retrieved from the following API endpoint:

https://pokeapi.co/api/v2/pokemon/squirtle

Write a function `pokemon_table(payload)` that accepts the parsed API response and returns a one-row pandas DataFrame containing these columns in order:

- `name`: the Pokémon’s name.
- `height`: its height in decimetres.
- `weight`: its weight in hectograms.

Extract the values from `payload`. Preserve the original units and do not hard-code the values. Store the DataFrame in `pokemon_df` before returning it.

```python
# Request real API data.
pokemon_data = get_json("https://pokeapi.co/api/v2/pokemon/squirtle")

...
def pokemon_table(payload):
    pokemon_df = ...
    return pokemon_df

# Display the table after defining the function.
display(pokemon_table(pokemon_data))
```

```python
grader.check("q2")
```

## Q3. GitHub API — Repository Languages (16 points)

Run the provided API request to retrieve five recently created public repositories.

Using `repository_data`:

1. Store the number of retrieved repositories in `repository_count`.
2. Create a pandas Series named `language_counts` showing the number of repositories for each language, sorted by count from highest to lowest.
3. Replace missing language values with `"Unknown"` before counting.

Display both results. Do not define a function.

```python
github_params = {
    "sort": "created", "direction": "desc", "per_page": 5, "page": 1
}
repository_data = get_json(
    "https://api.github.com/users/joelgrus/repos", github_params
)

...
repository_count = ...
languages = ...
language_counts = ...

print("Repositories retrieved:", repository_count)
display(language_counts)
```

```python
grader.check("q3")
```

## Q4. Weather API — Summarize Precipitation (22)

Run the provided request to collect Houghton-area weather data for October 2025.

Using `weather_data["daily"]`:

1. Create `weather_df` with columns `date`, `max_temperature`, and `precipitation`, in that order.
2. Convert dates to pandas datetime and measurements to numeric. Preserve missing values.
3. Store the number of days with precipitation greater than zero in `rainy_days`.
4. Store the total precipitation over available days in `total_precipitation`. If all precipitation values are missing, the total should remain missing.

Display the table and summary values. Do not define a function.

```python
weather_params = {
    "latitude": 47.12, "longitude": -88.57,
    "start_date": "2025-10-01", "end_date": "2025-10-31",
    "daily": "temperature_2m_max,precipitation_sum",
    "timezone": "America/Detroit"
}
weather_data = get_json(
    "https://archive-api.open-meteo.com/v1/archive", weather_params
)

...
weather_df = ...
weather_df["date"] = ...
...
rainy_days = ...
total_precipitation = ...

display(weather_df.head())
print("Units:", weather_data["daily_units"])
print("Days with precipitation:", rainy_days)
print("Total precipitation over available days:", total_precipitation)
```

```python
grader.check("q4")
```

## Q5. iTunes API — Sort E-books by Price (16 points)

Run the provided request to search for “The Expanse” e-books.

1. Create `books_df` from the records in `books_data["results"]`.
2. Include only `trackName`, `trackId`, and `price`, renamed to `TrackName`, `TrackID`, and `Price`.
3. Sort by price from lowest to highest, placing missing prices last.
4. Reset the index and display the first five rows.

Preserve missing values. Do not define a function.

```python
books_params = {"term": "The Expanse", "entity": "ebook", "limit": 50}
books_data = get_json("https://itunes.apple.com/search", books_params)

...
books_df = ...
books_df.columns = ...
books_df = ...

display(books_df.head())
```

```python
grader.check("q5")
```

<!-- BEGIN QUESTION -->

## Q6. Visualize real weather data (16 points)

Use `weather_df` from Q4 to plot daily maximum temperature against date. Add markers, a title identifying the location and month, and labeled axes. Read the temperature unit from `weather_data["daily_units"]`; do not replace missing values with zero.

In 2–3 sentences, report the highest observed maximum temperature and its date. Explain how missing values, if present, affect your interpretation.

**Rubric:** correct dates/values and missing-value handling (8); title, labels, units and markers (3); accurate interpretation based on retrieved data (5).

_Type your answer here, replacing this text._

```python
...
# Plot the real weather data retrieved in Q4.
...
```

<!-- END QUESTION -->

## 🎉 Congratulations! You Have Completed Lab 07

To complete your submission:

1. **Run all notebook cells in order** to ensure that all outputs are generated correctly.
2. **Save your notebook** before exporting.
3. Run the submission cell to generate the ZIP file.
4. Upload the generated ZIP file to the **Lab 07 assignment on Gradescope**.

⚠️ Important: If you do not run all cells sequentially before exporting, some outputs may be missing and could affect your autograder results.

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False)
```

Tests for q1 (points: [2, 2, 3])
```
>>> family_summary(family_text) == ('Grandma', 'Dad', 34)
True
```
```
>>> result = family_summary(family_text)
>>> isinstance(result, tuple)
True
```
```
>>> result = family_summary(family_text)
>>> len(result) == 3
True
```


Tests for q2 (points: [3, 4])
```
>>> df = pokemon_table({'name': 'squirtle', 'height': 5, 'weight': 90})
>>> list(df.columns) == ['name', 'height', 'weight']
True
```
```
>>> df = pokemon_table({'name': 'squirtle', 'height': 5, 'weight': 90})
>>> df.iloc[0].tolist() == ['squirtle', 5, 90]
True
```


Tests for q3 (points: [2, 3, 3])
```
>>> repository_count == len(repository_data)
True
```
```
>>> isinstance(language_counts, pd.Series)
True
```
```
>>> language_counts.sum() == repository_count
np.True_
```


Tests for q4 (points: [3, 3, 4, 4])
```
>>> isinstance(weather_df, pd.DataFrame)
True
```
```
>>> list(weather_df.columns) == ['date', 'max_temperature', 'precipitation']
True
```
```
>>> 0 <= rainy_days <= len(weather_df)
np.True_
```
```
>>> pd.isna(total_precipitation) or total_precipitation >= 0
np.True_
```


Tests for q5 (points: [2, 3, 3])
```
>>> isinstance(books_df, pd.DataFrame)
True
```
```
>>> list(books_df.columns) == ['TrackName', 'TrackID', 'Price']
True
```
```
>>> len(books_df) == len(books_data['results'])
True
```
