# Lab 3: Exploratory Data Analysis and Statistics

### Assignment Objectives
Upon successful completion of this lab, students will be able to:
- Gain experience using and writing Python functions
- Apply basic statistical functions for data analysis
- Perform exploratory data visualization to understand data distributions and relationships
- Conduct basic data cleaning tasks to prepare datasets for analysis


```python
#  Import libraries 
import pandas as pd
import numpy as np
import matplotlib as mpl
import matplotlib.pyplot as plt
%matplotlib inline
import seaborn as sns

import otter
grader = otter.Notebook()
```

### Data Set - Mushroom

You will be using the [Seconday Mushroom Dataset](https://archive.ics.uci.edu/dataset/848/secondary+mushroom+dataset) available on the [UCI Machine Learning Repository](https://archive.ics.uci.edu/)

Please use the mushroom data files in your "data" folder. 




### Exercise 1: Mushroom Class Balance

Load `data/secondary_data.csv` into a DataFrame named `data`.

Then:

1. Create `class_counts`, containing the number of observations in each mushroom class.
2. Create `poisonous_percent`, containing the percentage of mushrooms labeled poisonous (`p`), rounded to **two decimal places**.

Keep the original class labels (`e` and `p`) in `class_counts`.


```python
filepath = "data/secondary_data.csv"
data = ...

class_counts = ...
poisonous_percent = ...

class_counts, poisonous_percent
```

```python
grader.check("q1")
```

### Exercise 2: Data Exploration


You will want to explore if this data has any missing data.  And if it does, handle the missing values. 

#### Exercise 2a: Columns That Actually Contain Missing Data

Create a Series named `missing_columns` that contains **only columns with at least one missing value**.

Requirements:

- the index should contain the column names,
- the values should be the number of missing observations,
- sort the Series from the **largest** missing count to the smallest.


```python
missing_counts_all = ...
missing_columns = ...
missing_columns

```

```python
grader.check("q2a")
```

#### Exercise 2b: Complete Observations

An observation is considered **complete** when it has no missing values in any column.

Calculate:

- `complete_row_count`: the number of complete rows,
- `complete_row_percent`: the percentage of complete rows in the dataset, rounded to **two decimal places**.


```python
complete_row_count = ...
complete_row_percent = ...

complete_row_count, complete_row_percent

```

```python
grader.check("q2b")
```

The amount of missing data precludes eliminating the missing data, therefore additional analysis steps should understand there may be missing values. 

#### Exercise 2c: Fill One Categorical Variable

Create a copy of `data` named `mushroom_filled`.

In `mushroom_filled`, replace missing values in `cap-surface` with the string `"unknown"`.

Then store the number of remaining missing values in `cap-surface` in `cap_surface_missing_after`.

Do **not** modify the original `data` DataFrame.


```python
mushroom_filled = ...
mushroom_filled['cap-surface'] = ...
cap_surface_missing_after = ...

cap_surface_missing_after
```

```python
grader.check("q2c")
```

### Exercise 3: Data Exploration and Visualization



#### Exercise 3a: Spread of Cap Diameter

For the `cap-diameter` variable, calculate:

- `cap_range`: maximum minus minimum,
- `cap_iqr`: third quartile minus first quartile,
- `cap_cv`: coefficient of variation, calculated as standard deviation divided by mean.

Use Pandas methods for the calculations.


```python
cap_range = ...
cap_iqr = ...
cap_cv = ...

cap_range, cap_iqr, cap_cv

```

```python
grader.check("q3a")
```

#### Exercise 3b: Relationship Between Cap Diameter and Stem Width

Create a function `plot_cap_stem_relationship(df)` that:

- calculates the Pearson correlation between `cap-diameter` and `stem-width`,
- creates a scatter plot with `cap-diameter` on the x-axis and `stem-width` on the y-axis,
- includes a title and axis labels,
- returns the correlation value and the Matplotlib `Axes` object.


```python
def plot_cap_stem_relationship(df):
    corr = ...
    ax = ...

    # Add a title and axis labels

    return corr, ax

cap_stem_corr, cap_stem_ax = plot_cap_stem_relationship(data)

```

```python
grader.check("q3b")
```

#### Exercise 3c: Most Common Cap Colors

Create a function `plot_top_cap_colors(df, n=5)` that:

- finds the `n` most common values in `cap-color`,
- stores their counts in descending order,
- creates a **horizontal bar plot** of those counts,
- returns the count Series and the Matplotlib `Axes` object.

Use the default value `n=5` when calling the function.


```python
def plot_top_cap_colors(df, n=5):
    top_colors = ...
    ax = ...

    # Add a title and axis labels

    return top_colors, ax

top_cap_colors, cap_color_ax = plot_top_cap_colors(data)

```

```python
grader.check("q3c")
```

#### Exercise 3d: Mushroom Class by Season

Create a function `plot_class_by_season(df)` that:

- builds a cross-tabulation of `season` versus `class`,
- creates a **stacked bar chart** from the cross-tabulation,
- includes a title and axis labels,
- returns the cross-tabulation DataFrame and the Matplotlib `Axes` object.


```python
def plot_class_by_season(df):
    season_class_table = ...
    ax = ...

    # Add a title and axis labels

    return season_class_table, ax

season_class_table, season_class_ax = plot_class_by_season(data)
```

```python
grader.check("q3d")
```

## Data Set - Auto-MPG 

For the remaining problems, you will using the Auto-MPG data set available in the UCI ML Repository. 

https://archive.ics.uci.edu/ml/datasets/Auto+MPG

### Exercise 4: Auto-MPG Coverage

Load the Auto-MPG dataset into `data` using the column names provided below.

After loading the data:

1. Store the earliest model year in `earliest_model_year`.
2. Store the latest model year in `latest_model_year`.
3. Create `cylinder_counts`, containing the number of vehicles for each cylinder value, sorted by cylinder number.

Use the supplied Auto-MPG URL and column names.


```python
url = "https://archive.ics.uci.edu/ml/machine-learning-databases/auto-mpg/auto-mpg.data"
cNames = ['mpg', 'cyl', 'disp', 'hp', 'weight',
          'accel', 'model year', 'origin', 'car name']

data = ...

earliest_model_year = ...
latest_model_year = ...
cylinder_counts = ...

earliest_model_year, latest_model_year, cylinder_counts
```

```python
grader.check("q4")
```

### Exercise 5: Data Exploration

#### Exercise 5a: Locate Cars With Unknown Horsepower

The `hp` column uses `?` to mark unknown horsepower.

Create:

- `unknown_hp_rows`: a DataFrame containing only rows where `hp` is `?`,
- `unknown_hp_cars`: a Python list containing the corresponding values from `car name`.


```python
unknown_hp_rows = ...
unknown_hp_cars = ...

unknown_hp_rows[['hp', 'car name']]
```

```python
grader.check("q5a")
```

#### Exercise 5b: Vehicle Origins

Create `origin_counts`, containing the number of vehicles from each numeric origin code.

Then determine the most common origin code and store it in `most_common_origin`.

Sort `origin_counts` by origin code.


```python
origin_counts = ...
most_common_origin = ...

origin_counts, most_common_origin
```

```python
grader.check("q5b")
```

**What is wrong with this information?** 

Horsepower, `hp` is considered an object.  But, this should be a float.  What is the problem?

```python
data[data['hp'] == '?']
```

There are some missing values in this column that need to be handled.

We could change our import function to recognize "?" values as `NaN`.  Here let's instead create a function that replaces the "?" values after the data is imported.  



#### Exercise 5c: Median-Impute Horsepower

Create a function `CleanHorsepower(df)` that cleans the `hp` column.

The function should:

- work on a copy of the input DataFrame,
- convert `?` to missing values,
- convert `hp` to numeric,
- calculate the median of the available horsepower values,
- replace missing horsepower values with that median,
- return both the cleaned DataFrame and the median used.

Apply the function and store the outputs in `data` and `hp_median`.


```python
def CleanHorsepower(df):
    cleaned = ...

    # Clean and median-impute the hp column

    hp_median = ...
    return cleaned, hp_median

data, hp_median = CleanHorsepower(data)

data['hp'].head(), hp_median

```

```python
grader.check("q5c")
```

#### Exercise 5d: Build Analysis-Friendly Labels

Create a function `PrepareAutoLabels(df)` that prepares labels for later analysis.

The function should:

1. Standardize the existing column names to **title case** and replace spaces with `_`.
2. Add a new column named `Manufacturer` containing the **first word** of each `Car_Name`.
3. Add a new column named `Region` using this mapping:
   - `1` → `USA`
   - `2` → `Europe`
   - `3` → `Asia`
4. Return the resulting DataFrame.

Apply the function to `data`.


```python
def PrepareAutoLabels(df):
    prepared = ...

    # Standardize the existing column names
    # Add Manufacturer
    # Add Region

    return prepared

data = PrepareAutoLabels(data)

data.head()
```

```python
grader.check("q5d")
```

### Exercise 6: Data Exploration and Visualization 

Let's explore the data with some visualizations. 

#### Exercise 6a: Fuel Economy by Cylinder Count

Create a function `plot_mpg_by_cylinders(df)` that:

- calculates the mean `Mpg` for each value of `Cyl`,
- sorts the result by cylinder count,
- creates a bar plot of the mean MPG values,
- includes a title and axis labels,
- returns both the grouped Series and the Matplotlib `Axes` object.


```python
def plot_mpg_by_cylinders(df):
    mpg_by_cyl = ...
    ax = ...

    # Add a title and axis labels

    return mpg_by_cyl, ax

mpg_by_cyl, mpg_cyl_ax = plot_mpg_by_cylinders(data)

```

```python
grader.check("q6a")
```

#### Exercise 6b: Weight and Fuel Economy

Create a function `plot_weight_vs_mpg(df)` that:

- calculates the Pearson correlation between `Weight` and `Mpg`,
- creates a scatter plot with `Weight` on the x-axis and `Mpg` on the y-axis,
- includes a title and axis labels,
- returns the correlation value and the Matplotlib `Axes` object.


```python
def plot_weight_vs_mpg(df):
    weight_mpg_corr = ...
    ax = ...

    # Add a title and axis labels

    return weight_mpg_corr, ax

weight_mpg_corr, weight_mpg_ax = plot_weight_vs_mpg(data)
```

```python
grader.check("q6b")
```

<!-- BEGIN QUESTION -->

#### Exercise Bonus: Data Exploration and Analysis 

Create a function `plot_mpg_year_region()` that reports the mpg over the both groupings `Model Year` and `Region`.  Think about what is the appropriate type of plot.  

The function should: 

* take as input the DataFrame 
* report mpg over both groups
* make sure the plot looks good and has labels, title, etc. 

```python
def plot_mpg_year_region(df):
    
    
    return 


plot_mpg_year_region(data)
```

<!-- END QUESTION -->

## ✅ Congratulations! You Have Completed **Lab 03**

### 📤 Submission Instructions

Below, you will find a submission cell. **Running this cell will automatically generate a ZIP file** containing your autograded responses and submit it to **Gradescope** for the **Lab 03** assignment.

**Important notes:**
- ✔️ Verify that **all responses, outputs, and figures** required by each question are correctly generated and included in the submission.
- ✔️ Confirm that your submission appears under **Lab 03 – Assignment** on Gradescope after the process completes.
- ⚠️ If you encounter any issues during submission, **contact the course staff well before the deadline**.

Before running the submission cell:
- 🔁 Make sure you have **run all notebook cells in order** from top to bottom.
- 🖼️ Ensure that **all plots, images, and tables** are visible in the output.
- 💾 **Save your notebook** before exporting and submitting.

Once everything is verified, proceed to run the submission cell below.


## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False)
```

Tests for q1 (points: [1, 1, 1])
```
>>> isinstance(data, pd.DataFrame)
True
```
```
>>> isinstance(class_counts, pd.Series)
True
```
```
>>> set(class_counts.index) == {'e', 'p'}
True
```


Tests for q2a (points: [2, 2])
```
>>> isinstance(missing_columns, pd.Series)
True
```
```
>>> missing_columns.is_monotonic_decreasing
True
```


Tests for q2b (points: [1, 1])
```
>>> 0 <= complete_row_count <= len(data)
True
```
```
>>> 0 <= complete_row_percent <= 100
True
```


Tests for q2c (points: [1, 1])
```
>>> mushroom_filled is not data
True
```
```
>>> cap_surface_missing_after == 0
np.True_
```


Tests for q3a (points: [1, 1, 1])
```
>>> cap_range > 0
np.True_
```
```
>>> cap_iqr > 0
np.True_
```
```
>>> cap_cv > 0
np.True_
```


Tests for q3b (points: [1])
```
>>> -1 <= cap_stem_corr <= 1
np.True_
```


Tests for q3c (points: [1])
```
>>> isinstance(top_cap_colors, pd.Series) and isinstance(cap_color_ax, mpl.axes.Axes)
True
```


Tests for q3d (points: [1, 1])
```
>>> isinstance(season_class_table, pd.DataFrame) and isinstance(season_class_ax, mpl.axes.Axes)
True
```
```
>>> season_class_table.to_numpy().sum() == len(data)
np.True_
```


Tests for q4 (points: [1, 1, 2])
```
>>> isinstance(data, pd.DataFrame)
True
```
```
>>> earliest_model_year <= latest_model_year
np.True_
```
```
>>> cylinder_counts.index.is_monotonic_increasing
True
```


Tests for q5a (points: [1, 1, 1])
```
>>> isinstance(unknown_hp_rows, pd.DataFrame)
True
```
```
>>> isinstance(unknown_hp_cars, list)
True
```
```
>>> (unknown_hp_rows['hp'] == '?').all()
np.True_
```


Tests for q5b (points: [1, 1])
```
>>> origin_counts.sum() == len(data)
np.True_
```
```
>>> most_common_origin in origin_counts.index
True
```


Tests for q5c (points: [1, 1])
```
>>> pd.api.types.is_numeric_dtype(data['hp'])
True
```
```
>>> data['hp'].isna().sum() == 0
np.True_
```


Tests for q5d (points: [1, 1])
```
>>> 'Manufacturer' in data.columns and 'Region' in data.columns
True
```
```
>>> set(data['Region'].dropna().unique()) == {'USA', 'Europe', 'Asia'}
True
```


Tests for q6a (points: [1])
```
>>> isinstance(mpg_by_cyl, pd.Series) and isinstance(mpg_cyl_ax, mpl.axes.Axes)
True
```


Tests for q6b (points: [1])
```
>>> -1 <= weight_mpg_corr <= 1
np.True_
```
